```python
# Initialize Otter
import otter
grader = otter.Notebook("ps2.ipynb")
```

# STATS 507
## Problem Set 2
All functions will be tested by visible as well as hidden tests. The maximum amount of time any function is allowed to run is 10 seconds. If your program runs beyond 10 secs, the tests fail.


### Question 1: Fun with Strings (3 points)

In this problem, you'll implement a few simple functions for dealing
with strings. You need not perform any error checking in any of the
functions for this problem.

**1(a)** (1 pts.) Let us say that a word is "abecedarian" if its letters appear in
    alphabetical order (repeated letters are okay). So, for example,
    "adder" and "beet" are abecedarian, whereas "dog" and "cat" are not.
    Write a function `is_abecedarian`, which takes a single argument in the form of a
    string and returns `True` if the argument is abecedarian and `False` otherwise.
    Here you may assume that the input consists only of alphabetic
    characters and spaces. Your function should ignore spaces, so that
    the string `abcd efgh xyz` is considered abecedarian. Your function should ignore
    capitalization, so that `Abc`, `A bc` and `aA Bc` are all considered
    abecedarian.
    
Hint: Take a look at the [zip](https://docs.python.org/3/library/functions.html#zip) function that might help you make your program a little simple

```python
import collections # using collections module is not required but you may use it if you want
```

```python
def is_abecedarian(string):
    """
    Returns whether a string is abecedarian or not.
    """
    s = "".join(string.lower().split())
    # Create pairs of consecutive characters
    pairs = list(zip(s, s[1:]))
    # Check if each pair is in order
    for i, j in pairs:
        if i > j:
            return False
    return True 
```

```python
#is_abecedarian("aba")
#isinstance(is_abecedarian('abC'), bool)
#is_abecedarian("dog")
#is_abecedarian("cat")
#is_abecedarian("adder")
#'a' > 'b'
#is_abecedarian('Abc')
#is_abecedarian('a b  c')
```

```python
grader.check("q1a")
```

Output:
```
q1a results: All test cases passed!
```

**1(b)** (2 pts.) Write a function called `double_vowels` that takes a string as its only argument
    and returns that string with all the vowels duplicated. For the
    purposes of this question, the vowels are the letters "a e i o u".
    Thus, `double_vowels(’cat’)` should return `caat`, `double_vowels(’audio’)` should return `aauudiioo`, `double_vowels(’Aa’)` should return `AAaa`,
    etc. 
    

```python
def double_vowels(string):
    """
    Returns a string with all the vowels duplicated
    """
    dbl_string = ''
    for s in string:
        if s.lower() in 'aeiou':
            dbl_string += s*2
        else:
            dbl_string += s
    return dbl_string
        
        
```

```python
#double_vowels('About')
#double_vowels('francisco')
```

Output:
```
'fraanciiscoo'
```

```python
grader.check("q1b")
```

Output:
```
q1b results: All test cases passed!
```

### Question 2: Fun with Lists (4 points)

In this problem, you'll implement a few very simple list operations.

**2(a)** (1 pts.) Write a function `list_reverse` that takes a list as an argument and returns
    that list, reversed. That is, given the list `[1,2,3]`, your function
    should return the reversed list, `[3,2,1]`. Your function should raise an
    appropriate error in the event that the input is not a list.

```python
def list_reverse(lis):
    '''
    Returns the reversed list. 
    '''
    if not isinstance(lis, list):
        raise TypeError('Input must be a list!')
    else:
        return lis[::-1]
```

```python
#a = ['a', 'b']
#a[::-1]
#list_reverse(['12', 4])
#list_reverse('aeiou')
```

```python
grader.check("q2a")
```

Output:
```
q2a results: All test cases passed!
```

**2(b)** (1 pts.) Write a function `is_sorted` that takes a sequence `seq` as its only argument
    and returns `True`if the sequence is sorted in non-decreasing order and
    returns `False` otherwise. You may assume that `seq` is, in fact, a Python
    sequence. Your function should require a single traversal of the
    list. You may assume that all
    elements in the input sequence `seq` support the comparison operations
    (`==`, `<`/`>`, `>=`, etc), so that there is no need for error checking.
    (Indeed, if you try to make the comparison, say, `1 < ’cat’`, Python will
    raise an error for you, anyway.) **Note:** this problem illustrates
    a particularly useful aspect of Python's dynamic typing. It is
    possible to write this function while being agnostic as to the type
    of the input variable. It requires only that `seq` supports indexing
    and that the elements in `seq` support the comparison operations.

```python
def is_sorted(seq):
    """
    Returns whether the sequence is sorted.
    """
    return sorted(seq) == seq
```

```python
#is_sorted([1,2,3])
#is_sorted([4,44,1])
# also works for letters
#is_sorted(['c', 'd'])
#is_sorted(['b', 'a'])
```

```python
grader.check("q2b")
```

Output:
```
q2b results: All test cases passed!
```

**2(c)** (2 pts.) This one is a common coding interview question. Write a function
    called `binary_search` that takes two arguments, a list of integers `t` (which is
    guaranteed to be sorted in ascending order) and an integer `elmt`, and
    returns `True` if `elmt` appears in list `t` and `False` otherwise. Of course, you
    could do this with the `in` operator, but that will be slow when the
    list is long. Instead, you
    should use *binary search* which brings down the time complexity to **[O(log n)](https://stackoverflow.com/questions/1592649/examples-of-algorithms-which-has-o1-on-log-n-and-olog-n-complexities)**: To look for `elmt`, first look at the
    "middle" element of the list `t`. If it's a match, return `True`. If it
    isn't a match, compare `elmt` against the "middle" element, and recurse,
    searching the first or second half of the list depending on whether
    `elmt` is bigger or smaller than the middle element. 
    
**Hint:** be careful of the *base cases*: What should you do when `t` is empty,
    length 1, length 2, etc.? 
    
**Note:** your solution must actually make
    use of binary search to receive credit, and your solution must not
    use any built-in sorting or searching functions. 


```python
def binary_search(t, elmt):
    """
    Returns whether elmt is in t
    """
    if len(t) == 0:
        return False
    else:
        mid = len(t) // 2
        if t[mid] == elmt:
            return True
        else:
            if elmt < t[mid]:
                return binary_search(t[:mid], elmt)
            else:
                return binary_search(t[mid+1:], elmt)
```

```python
#binary_search([1,2,3,4], 2)
#binary_search([11,12,13,14,15], 10)
#binary_search([], 2)
```

```python
grader.check("q2c")
```

Output:
```
q2c results: All test cases passed!
```

### Question 3: More Fun with Strings (4 points)

In this problem, you'll implement some very simple counting operations
that are common in fields like biostatistics and natural language
processing. You need not perform any error checking in the functions for
this problem.

**3(a)** (2 pts.) Write a function called `char_hist` that takes a string as its argument and
    returns a dictionary whose keys are characters and values are the
    number of times each character appeared in the input. So, for
    example, given the string "gattaca", your function should return a
    dictionary with key-value pairs `g:1, a:3, t:2, c:1`. Your function should count *all*
    characters in the input (including spaces, tabs, numbers, etc). The
    dictionary returned by your function should have as its keys all and
    only the characters that appeared in the input (i.e., you don't need
    to have a bunch of keys with value 0). Your function should count
    capital and lower-case letters as the same, and key on the
    lower-case version of the character, so that `G` and `g` are both
    counted as the same character, and the corresponding key in the
    dictionary is `g`.

```python
def char_hist(string):
    """
    Returns a dictionary whose keys are characters and values
    are the number of times each character appeared in the input.
    """
    char_dict = {}
    for char in string.lower():
        if char.isalnum():
            if char in char_dict:
                char_dict[char] += 1
            else:
                char_dict[char] = 1
        elif char in char_dict:
            char_dict[char] += 1
        else:
            char_dict[char] = 1
    return char_dict
        
        
```

```python
#str = 'google'
#char_hist(str)3
#str = '2024'
#char_hist(str)
#char_hist('frank brady')
# char_hist('the moose is loose')
# char_hist('two + two')
#str = 'frank + ==  brady'
#char_hist(str)
#char_hist('gattaca')
```

```python
grader.check("q3a")
```

Output:
```
q3a results: All test cases passed!
```

**3(b)** (2 pts.) In natural language processing and bioinformatics, we often want to
    count how often characters or groups of characters appear. Pairs of
    words or characters are called "bigrams". For our purposes in this
    problem, a bigram is a pair of characters. As an example, the string
    `mississippi` contains the following bigrams, in order: {'mi', 'is', 'ss', 'si', 'is', 'ss', 'si', 'ip', 'pp', 'pi'}. Write a function called `bigram_hist` that takes a string as its
    argument and returns a dictionary whose keys are 2-tuples of
    characters and values are the number of times that pair of
    characters appeared in the string. So, for example, when called on
    the string `mississippi`, your function should return a dictionary with keys
    {(m,i), (i,s), (s,s), (s,i), (i,p), (i,p), (p,p), (p,i)}
    and respective count values 1,2,2,2,1,1,1. As another example,
    if the two-character string `ab` occurred four times in the input,
    then your function should return a dictionary that includes the
    key-value pair `(a,b):4`. Your function should handle all characters
    (alphanumerics, spaces, punctuation, etc). So, for example, the
    string `cat, dog` includes the bigrams `t,`, `, ` and ` d`. As in the previous
    subproblem, the dictionary produced by your function should only
    include pairs that actually appeared in the input, so that the
    absence of a given key implies that the corresponding two-character
    string did not appear in the input. Also as in the previous
    subproblem, you should count upper- and lower-case letters as the
    same, so that `GA`, `Ga`, `gA` and `ga` all count for the same key, ``.

```python
def bigram_hist(string):
    """
    Returns a dictionary whose keys are 2-tuples of characters
    and values are the number of times that pair of characters
    appeared in the string.
    """
    big_dict = {}
    for i in range(1, len(string)):
        bigram = (string[i-1].lower(), string[i].lower())
        if bigram in big_dict:
            big_dict[bigram] += 1
        else:
            big_dict[bigram] = 1
    return big_dict
```

```python
#bigram_hist('moo')
#bigram_hist('ca,ca,cat')
#bigram_hist("&&&&&&&&")
#bigram_hist('cat, dog')
#bigram_hist('CAt, cat, dog')
```

```python
grader.check("q3b")
```

Output:
```
q3b results: All test cases passed!
```

### Question 4: Tuples as Vectors (5 points)

In this problem, we’ll see how we can use tuples to represent vectors. Later in the semester, we’ll see the Python numpy and scipy packages, which provide objects specifically meant to enable matrix and vector operations, but for now tuples are all we have. So, for this problem we will represent a d-dimensional vector by a length-d tuple of floats.

**4(a)** (1 pts.) Implement a function called `vec_scalar_mult`, which takes two arguments: a tuple of numbers (floats and/or integers) `t` and a number (float or integer) `s` and returns a tuple of the same length as `t`, with its entries equal to the entries of `t` multiplied by `s`. That is, `vec_scalar_mult` implements multiplication of a vector by a scalar. Your function should check to make sure that the types of the input are appropriate (e.g., that `s` is a float or integer), and raise a TypeError with a suitable error message if the types are incorrect. However, your function should gracefully handle the case where the input `s` is an integer rather than a float, or the case where some or all of the entries of the input tuple are integers rather than floats. Hint: you may find it useful for this subproblem and the next few that follow it to implement a function that checks whether or not a given tuple is a “valid” vector (i.e., checks if a variable is a tuple and checks that its entries are all floats and/or integers).

```python
def valid_vector(vec):
    if isinstance(vec, tuple):
        # Check if all elements in the tuple are ints or floats
        return all(isinstance(v, (int, float)) for v in vec)   
    else:
        return False
```

```python
def vec_scalar_mult(t, s):
    """
    Returns multiplication of a vector by a scalar.
    """
    if not isinstance(s, (int, float)):
        raise TypeError('s must be int or float')
    if not valid_vector(t):
        raise TypeError("t should be a tuple of float(s) and/or integer(s)")
    return tuple(v * s for v in t)
```

```python
#vec_scalar_mult((1,2), 2)
#vec_scalar_mult((1,1,1.3), 2)
#vec_scalar_mult((1,1,1), 5.6)
```

```python
grader.check("q4a")
```

Output:
```
q4a results: All test cases passed!
```

**4(b)** (1 pts.) Implement a function called `vec_inner_product` which takes two "vectors" (i.e.,
    tuples of floats and/or ints) as its inputs and outputs a float
    corresponding to the inner product of these two vectors. Recall that
    the inner product of vectors $x,y \in \mathbb{R}^d$ is given by
    $\sum_{j=1}^d x_j y_j$. Your function should check whether or not
    the two inputs are of the correct type (i.e., both tuples), and
    raise a `TypeError` if not. Your function should also check whether or not
    the two inputs agree in their dimension (i.e., length, so that the
    inner product is well-defined), and raise a `ValueError` if not.

```python
def vec_inner_product(a, b):
    """
    Returns the inner product of two vectors.
    """
    if not valid_vector(a) or not valid_vector(b):
        raise TypeError('both a and b must be tuples of int or float')
    if len(a) != len(b):
        raise ValueError('a and b must agree in dimension')
    else:
        return sum([xj*yj for xj,yj in zip(a,b)])
```

```python
#vec_inner_product((1,1), (2,3.3))
#vec_inner_product([1,2,3], (2,3))
#vec_inner_product({1:2}, (1,2))
```

```python
grader.check("q4b")
```

Output:
```
q4b results: All test cases passed!
```

**4(c)** (1 pts.) It is natural, following the above, to extend our scheme to the case
    of matrices. Recall that a matrix is simply a box of numbers. If you
    are not already familiar with matrices, feel free to look them up on
    Wikipedia or in any linear algebra textbook. We will represent a
    matrix as a *tuple of tuples*, i.e., a tuple whose entries are
    themselves tuples. We will represent an $m$-by-$n$ matrix as an
    $m$-tuple of $n$-tuples. To be more concrete, suppose that we are
    representing an $m$-by-$n$ matrix $M$ as a variable `my_mx`. Then `my_mx` will
    be a length-$m$ tuple of $n$-tuples, so that the $i$-th row of the
    matrix is given (as a vector) by the $i$-th entry of tuple `my_mx`.

Write a function `is_valid_matrix` that takes a single argument and returns a
    Boolean, which is `True` if the given argument is a tuple that validly
    represents a matrix as described above, and returns `False` otherwise. A
    valid matrix will be a tuple of tuples such that

-   Every element of the tuple is itself a tuple,

-   each of these tuples is the same length, and

-   every element of each of these tuples is a number (i.e., a float or integer).

```python
def all_tuple(mat):
    return all(isinstance(tpl, tuple) for tpl in mat)

def same_len(mat):
    return all(len(tpl) == len(mat[0]) for tpl in mat)

def all_num(mat):
    return all(isinstance(i, (float,int))  for tpl in mat for i in tpl)

def is_valid_mat(mat):
    """
    Returns whether mat is a valid matrix.
    """
    # each test defined as a separate function
    if  all_tuple(mat) and same_len(mat) and all_num(mat):
        return True
    else:
        return False
```

```python
#tpl = ((1,0),(0,1))
#tpl = (('1','0'), (0,1))
#is_valid_mat(tpl)
#mat = ((1,0),(0,1))
#mat[0]
#len(mat[0])
#mat = ((1,0),(0,1))
#is_valid_mat(mat)
#is_valid_mat([1,2])
```

```python
grader.check("q4c")
```

Output:
```
q4c results: All test cases passed!
```

**4(d)** (2 pts.) Write a function `mat_vec_mult` that takes a matrix (i.e., tuple of tuples) and
    a vector (i.e., a tuple) as its arguments, and returns a vector
    (i.e., a tuple of numbers) that is the result of multiplying the
    given vector by the given matrix (we are treating vectors as column
    vectors here, so the matrix acts "on the left"). Again, if you are
    not familiar with matrix-vector multiplication, refer to Wikipedia
    or any linear algebra textbook. Your function should check that all
    the supplied arguments are reasonable (e.g., using your function
    `is_valid_mat`), and raise an appropriate error if not. **Hint:** you may find
    it useful to make use of the inner-product function that you defined
    previously.

```python
def mat_vec_mult(mat, vec):
    """
    Returns the result of multiplying the given vector by the given matrix.
    """
    if not valid_vector(vec):
        raise TypeError("Vec must be tuple of int or float")
    if not is_valid_mat(mat):
        raise TypeError("Mat must be a list of tuples")
    if len(mat[0]) != len(vec):
        raise ValueError("Columns of mat must == rows of vec")
    return tuple(vec_inner_product(row, vec) for row in mat)
```

```python
mat = ((1,1),(1,2))
vec = (2,2)
mat_vec_mult(mat, vec)

```

Output:
```
(4, 6)
```

```python
grader.check("q4d")
```

Output:
```
q4d results: All test cases passed!
```

### Question 5: More Fun with Vectors (4 points)

In the previous problem, you implemented matrix and vector operations
using tuples to represent vectors. In many applications, it is common to
have vectors of dimension in the thousands or millions, but in which
only a small fraction of the entries are nonzero. Such vectors are
called *sparse* vectors, and if we tried to represent them as tuples, we
would be using thousands of entries just to store zeros, which would
quickly get out of hand if we needed to store hundreds or thousands of
such vectors.

A reasonable solution is to instead represent a sparse vector (or
matrix) by only storing its non-zero entries, with (index, value) pairs.
We will take this approach in this problem, and represent vectors as
dictionaries with positive integer keys (so we index into our vectors
starting from 1, just like in MATLAB and R). A *valid* sparse vector
will be a dictionary that has the properties that (1) all its indices
are positive integers, and (2) all its values are floats.

**5(a)** (2 pts.) Write a function `is_sparse_vector` that takes one argument, and returns `True` if and
    only if the input is a valid sparse vector, and returns `False`
    otherwise. **Note:** your function should *not* assume that the
    input is a dictionary.

```python
def is_sparse_vector(d):
    """
    Returns whether the given variable represents a sparse vector
    """
    if not isinstance(d, dict):
        return False
    for key, value in d.items():
        if not isinstance(key, int) or key <= 0:
            return False
        if not isinstance(value, float):
            return False
    return True
```

```python
grader.check("q5a")
```

Output:
```
q5a results: All test cases passed!
```

**5(b)** (2 pts.) Write a function `inner_sparse` that takes two "sparse vectors" as its inputs,
    and returns a float that is the value of the inner product of the
    vectors that the inputs represent. Your function should raise an
    appropriate error in the event that either of the inputs is not a
    valid sparse vector. Note that by our definition, a sparse vector
    has no specified dimension, so there is no need to check that the
    dimensions of the arguments agree.

```python
def inner_sparse(vec1, vec2):
    """
    Returns the inner product of two sparse vectors
    """
    if not is_sparse_vector(vec1) or not is_sparse_vector(vec2):
        raise TypeError("Both inputs must be sparse vectors")
    result = 0.0
    for key, value in vec1.items():
        if key in vec2:
            result += value * vec2[key]
    return result
```

```python
grader.check("q5b")
```

Output:
```
q5b results: All test cases passed!
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Upload this .zip file to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Output:
```
<IPython.core.display.HTML object>
```

Tests for q1a (points: 1)
```
>>> assert isinstance(is_abecedarian("aba"), bool)
>>> assert not is_abecedarian("dog")
>>> assert not is_abecedarian("cat")
>>> assert is_abecedarian("adder")
>>> assert is_abecedarian("beet")
```
```
>>> assert is_abecedarian("Abc")
>>> assert is_abecedarian("A bc")
>>> assert is_abecedarian("aA Bc")
```


Tests for q1b (points: 2)
```
>>> assert isinstance(double_vowels("aba"), str)
>>> assert double_vowels("cat") == "caat"
>>> assert double_vowels("audio") == "aauudiioo"
>>> assert double_vowels("Aa") == "AAaa"
```


Tests for q2a (points: 1)
```
>>> assert isinstance(list_reverse([1,2,3]), list)
>>> assert list_reverse([1,2,3]) == [3,2,1]
>>> assert list_reverse([1, 2, 'three', 4.0]) == [4.0, 'three', 2, 1]
>>> assert list_reverse([]) == []
```
```
>>> for invalid in ["12", 4]:
...     try:
...         list_reverse(invalid)
...     except TypeError:
...         pass
...     else:
...         raise Exception(f"{invalid} should raise an appropriate type of exception")
```


Tests for q2b (points: 1)
```
>>> assert is_sorted([1])
>>> assert is_sorted([1, 2, 3])
>>> assert is_sorted([1, 1, 1, 2, 3, 4, 4])
>>> assert is_sorted(["a", "a", "ab", "b"])
>>> assert not is_sorted([1, 2, 1])
>>> assert not is_sorted(["abb", "ab"])
```


Tests for q2c (points: 2)
```
>>> assert not binary_search([], 1)
>>> assert not binary_search([2], 1)
>>> assert not binary_search([2, 3], 1)
>>> assert not binary_search([-1, 0], 1)
>>> assert not binary_search([-1, 2], 1)
>>> assert not binary_search(list(range(1, 100, 2)), 98)
>>> assert not binary_search(list(range(1, 100, 2)), 102)
>>> assert not binary_search(list(range(1, 100, 2)), -1)
>>> assert not binary_search(list(range(1, 100, 2)), 50)
>>> assert binary_search([1], 1)
>>> assert binary_search([1, 3], 1)
>>> assert binary_search([-1, 1], 1)
>>> assert binary_search(list(range(1, 100, 2)), 99)
>>> assert binary_search(list(range(1, 100, 2)), 51)
>>> assert binary_search(list(range(1, 100, 2)), 1)
```


Tests for q3a (points: 2)
```
>>> assert isinstance(char_hist("abc"), dict)
>>> assert char_hist("gattaca") == {"g": 1, "a": 3, "t": 2, "c": 1}
>>> assert char_hist("GattAca   _+=") == {"g": 1, "a": 3, "t": 2, "c": 1, " ": 3, "_": 1, "+": 1, "=": 1}
```


Tests for q3b (points: 2)
```
>>> assert isinstance(bigram_hist("abc"), dict)
>>> assert bigram_hist("gattaca") == {("g", "a"): 1, ("a", "t"): 1, ("t", "t"): 1, ("t", "a"): 1, ("a", "c"): 1, ("c", "a"): 1}
>>> assert bigram_hist("mississippi") == {("m", "i"): 1, ("i", "s"): 2, ("s", "s"): 2, ("s", "i"): 2, ("i", "p"): 1, ("i", "p"): 1, ("p", "p"): 1, ("p", "i"): 1}
>>> assert bigram_hist("MisSiSsipPi") == bigram_hist("mississippi")
>>> assert bigram_hist("m+m+m") == {("m", "+"): 2, ("+", "m"): 2}
```


Tests for q4a (points: 1)
```
>>> assert isinstance(vec_scalar_mult((1,), 2), tuple)
>>> assert len(vec_scalar_mult((1, 2, 3), 3)) == 3
>>> assert vec_scalar_mult((1, 2, 3), 3) == (3, 6, 9)
>>> t = (1, 2, 3)
>>> s = 2.3
>>> res = vec_scalar_mult(t, s)
>>> for i in range(len(res)):
...     assert abs(res[i] - (i + 1) * 2.3) < 1e-10
```
```
>>> for i in [[1, 2, 3], 2, {1: 2}]:
...     for j in [2, 2.3]:
...         try:
...             vec_scalar_mult(i, j)
...         except TypeError:
...             pass
...         else:
...             raise Exception(f"vec_scalar_mult({i!r}, {j!r}) should raise an exception")
>>> for i in [(1, 2, 3), (1, 2.0, 3.0)]:
...     for j in ["j", (1, 2), [1, 2], {1: 2}]:
...         try:
...             vec_scalar_mult(i, j)
...         except TypeError:
...             pass
...         else:
...             raise Exception(f"vec_scalar_mult({i!r}, {j!r}) should raise an exception")
```


Tests for q4b (points: 1)
```
>>> assert isinstance(vec_inner_product((1, 2), (3, 4)), int)
>>> assert isinstance(vec_inner_product((1.0, 2), (3, 4)), float)
>>> assert vec_inner_product((1, 2, 3), (1, 2, 3)) == 14
>>> t = (1.6, 2.7, 3.9)
>>> s = (1.2, 2.3, -2.4)
>>> res = vec_inner_product(t, s)
>>> nres = sum([t[i] * s[i] for i in range(3)])
>>> assert abs(res - nres) < 1e-10
```
```
>>> for i in [[1, 2, 3], 2, {1: 2}]:
...     for j in [(2, 3), (2.3,)]:
...         try:
...             vec_inner_product(i, j)
...         except TypeError:
...             pass
...         else:
...             raise Exception(
...                 f"vec_inner_product({i!r}, {j!r}) should raise an exception"
...             )
>>> for i in [[1, 2, 3], 2, {1: 2}]:
...     for j in [(2, 3), (2.3,)]:
...         try:
...             vec_inner_product(j, i)
...         except TypeError:
...             pass
...         else:
...             raise Exception(
...                 f"vec_inner_product({i!r}, {j!r}) should raise an exception"
...             )
>>> 
>>> for i in [(1, 2, 3), (4, 5, 6.0), (7.0, 8.0, 9.0)]:
...     for j in [(1, 2, 3, 4), (5, 6, 7, 8.0), (9.0, 10.0, 11.0, 12.0)]:
...         try:
...             vec_inner_product(j, i)
...         except ValueError:
...             pass
...         else:
...             raise Exception(
...                 f"vec_inner_product({i!r}, {j!r}) should raise an exception"
...             )
```


Tests for q4c (points: 1)
```
>>> assert isinstance(is_valid_mat(((1, 2), (3, 4))), bool)
>>> assert isinstance(is_valid_mat([1, 2]), bool)
>>> assert not is_valid_mat([[1, 2], [3, 4]])
>>> assert not is_valid_mat(("a", 4))
>>> assert not is_valid_mat(((1, 2), 4))
>>> assert not is_valid_mat(((1, "4"), (3, 4)))
>>> assert not is_valid_mat(((1, 2, 3), (4, 5)))
>>> assert is_valid_mat(((1, 2, 3), (4, 5, 6), (7, 8, 9), (10, 11, 12)))
```


Tests for q4d (points: 2)
```
>>> assert isinstance(mat_vec_mult(((1,2), (3, 4)), (1,2)), tuple)
>>> mat = ((1.0,2.0), (3.0,4.0))
>>> vec = (1.0, 2.0)
>>> res = mat_vec_mult(mat, vec)
>>> resn = (5.0, 11.0)
>>> assert len(res) == 2
>>> for i in range(len(res)):
...     assert abs(res[i] - resn[i]) < 1e-10
```
```
>>> for i in [[[1,2], [3,4]], ('a', 4), (1, 2), ((1,2), 4), ((1,2,3), (4,5))]:
...     for j in [(2, 3), (2.3, )]:
...         try:
...             mat_vec_mult(i, j)
...         except TypeError:
...             pass
...         else:
...             raise Exception(f"mat_vec_mult({i!r}, {j!r}) should raise an exception")
>>> for i in [((1,2), (3,4))]:
...     for j in [[1,2], (1, 'c'), 'j', ((1,2), 3)]:
...         try:
...             mat_vec_mult(i, j)
...         except TypeError:
...             pass
...         else:
...             raise Exception(f"mat_vec_mult({i!r}, {j!r}) should raise an exception")
```


Tests for q5a (points: 2)
```
>>> assert isinstance(is_sparse_vector({1: 2}), bool)
>>> assert isinstance(is_sparse_vector([1, 2]), bool)
>>> for i in [3, (1, 2), [1, 2], {"a": 1.0}, {1: "a"}, {0: 2.0}]:
...     assert not is_sparse_vector(i)
>>> for i in [{1: 300.0, 2: 20.0, 400: 200.0}, {50: 200.0}]:
...     assert is_sparse_vector(i)
```


Tests for q5b (points: 2)
```
>>> assert isinstance(inner_sparse({1: 2.0}, {2: 3.0, 1: 3.0}), float)
>>> vec1 = {10: 2.0, 20: 3.0, 50: -1.0, 80: -20.0}
>>> vec2 = {300: 3.0, 20: 4.0, 50: 5.0, 9: 30.0}
>>> res = inner_sparse(vec1, vec2)
>>> resn = 7.0
>>> assert abs(res - resn) < 1e-8
```
```
>>> for i in [3, (1, 2), [1, 2], {"a": 1.0}, {1: "a"}, {0: 2.0}]:
...     for j in [{2: 3.0}, {2: 4.0}]:
...         try:
...             inner_sparse(i, j)
...         except TypeError:
...             pass
...         else:
...             raise Exception(f"inner_sparse({i!r}, {j!r}) should raise an exception")
```
